# SLU03 - Git Basics - Exercise notebook

```python
# Import for evaluating the results.
import hashlib
```

I hope you enjoyed the Learning Notebook! The Exercise Notebook is multiple choice so it's really important that you have a good grasp of everything before starting it. 

Always assign the letter of the correct option to the answer variable, e.g. `exercise_0 = 'A'`.

Good luck!

## Exercise 0: Stage your changes
This exercise is ungraded. It's purpose it to prepare the terrain for exercise 10.

- On the root of your Prep Course **workspace** repository, `ds-prep-workspace`, run the `git add .` command to stage everything.
- Run the `git status` command.
- Check the output, you should see that all files are in the staging area.

## Exercise 1: Version Control
What is the main feature of a VCS?

**A.** Tracking of changes to files in a given location.

**B.** Public sharing of files in a remote repository.

**C.** Uploading files to GitHub.

**D.** It contains a README file.

<!-- grade_id: cell-0ad663f9cf9def82, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_1.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_1 = 
```

<!-- grade_id: cell-28da65d304eb76f9, points: 2, grade: true, solution: false -->
```python
options = ["A", "B", "C", "D"]
assert isinstance(exercise_1,str), "The variable should be a string."
assert exercise_1.upper() in options, "Please choose one of these options: %s" % ', '.join(str(opt) for opt in options)
assert hashlib.sha256(bytes(exercise_1.upper(), encoding='utf8')).hexdigest() == \
"559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd", "Not correct, try again."
```

## Exercise 2: Git

Which one of the following statements is false: 

**A.** Git is a decentralized version control tool.

**B.** Users can clone entire repositories to their local systems.

**C.** Commits are possible even if offline.

**D.** Committing a file automatically uploads it to the remote repository.

<!-- grade_id: cell-71db91f1c18eebb3, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_2.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_2 =
```

<!-- grade_id: cell-ad5700b91f015868, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_2,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_2.upper(), encoding='utf8')).hexdigest() == \
"3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43", "Not correct."
```

## Exercise 3: Staging
Which of the following statements is true:

**A.** You can only have one file staged at a time.

**B.** You can only stage files after a commit.

**C.** The staging area is an intermediate storage for changes that will be part of the next commit.

**D.** All files are added to the staging area automatically.

<!-- grade_id: cell-e9c408d24e570684, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_3.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_3 =
```

<!-- grade_id: cell-6a4c40061538306a, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_3,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_3.upper(), encoding='utf8')).hexdigest() == \
"6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d", "Not correct."
```

## Exercise 4: Git add

Which command should you use when you want to stage all pdf files in the workspace?

**A.** `git add --all`

**B.** `git add .`

**C.** `git add *.pdf`

<!-- grade_id: cell-ebcecd4226f1dce6, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_4.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_4 =
```

<!-- grade_id: cell-df39aeb2d728fcd8, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_4,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_4.upper(), encoding='utf8')).hexdigest() == \
"6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d", "Not correct."
```

## Exercise 5: Git add and Git commit

What is the function of the commands `git add .`, `git commit`, and `git commit -a`?

**A.** The first stages all files in the repo, the second commits all staged files in the repo, and the last stages and commits all changed files.

**B.** The first stages all files that start with a dot, the second commits all staged files in the repo, and the last stages and commits all changed files.

**C.** The first stages all files in the repo, the second commits all modified files in the repo, and the last stages and commits all changed files.

<!-- grade_id: cell-c696448a0429c1f8, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_5.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_5 =
```

<!-- grade_id: cell-aff7a12614986551, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_5,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_5.upper(), encoding='utf8')).hexdigest() == \
"559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd", "Not correct."
```

## Exercise 6: Git status

In a workspace with five tracked files and a clean staging area, a new file is created. Additionally, one of the tracked files is changed. What would the `git status` command show?

**A.** Staged: 1, Modified: 1, Untracked: 0

**B.** Staged: 0, Modified: 1, Untracked: 1

**C.** Staged: 5, Modified: 1, Untracked: 1

**D.** Staged: 5, Modified: 0, Untracked: 1

<!-- grade_id: cell-ad3eb1b9f7256b3d, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_6.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_6 =
```

<!-- grade_id: cell-c266b16ef62872b4, points: 3, grade: true, solution: false -->
```python
assert isinstance(exercise_6,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_6.upper(), encoding='utf8')).hexdigest() == \
"df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c", "Not correct."
```

## Exercise 7: .gitignore

Which one of the following statements is true: 

**A.** You can stage a file that is listed in .gitignore if you add it explicitly with the `git add` command.

**B.** You can't ignore the .gitignore file itself.

**C.** You can use patterns, such as wildcards, in the .gitignore.

<!-- grade_id: cell-ce41a1f975fae767, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_7.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_7 =
```

<!-- grade_id: cell-3dc6db0d118d1a92, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_7,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_7.upper(), encoding='utf8')).hexdigest() == \
'6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d', "Not correct."
```

## Exercise 8: git push

What would happen if you tried to push changes to a remote repository that is ahead on commits compared to your local repository?

**A.** The remote repository will handle everything like magic.

**B.** The remote repository will throw an error because versions don't match.

**C.** The remote repository will throw an error because the remote repository is ahead on commits.

<!-- grade_id: cell-d8e48506ae23cd8e, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_8.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_8 =
```

<!-- grade_id: cell-808483824d7e116a, points: 3, grade: true, solution: false -->
```python
assert isinstance(exercise_8,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_8.upper(), encoding='utf8')).hexdigest() == \
"6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d", "Not correct."
```

## Exercise 9: Make commits useful

You have been working in your workspace for some time and you have developed code for two new features in two different folders, `feature_a/` and `feature_b/`. While testing the new features, you have created multiple output files in the `temp\` folder. You don't want to push these output files to the remote repo, but you also don't want to include them in the `.gitignore` because they are temporary. What is the best way to commit these changes?

**A.** `git commit -a -m "Add features a & b"`

**B.** `git add feature_a feature_b && git commit -m "Add features a & b"`

**C.** `git add feature_a && git commit -m "Add feature a"` and then `git add feature_b && git commit -m "Add feature b"`

<!-- grade_id: cell-d476eee500a7b97e, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_9.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_9 =
```

<!-- grade_id: cell-785cbbadef9d4f49, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_9,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_9.upper(), encoding='utf8')).hexdigest() == \
"6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d", "Not correct."
```

## Exercise 10: Git Workflow
This exercise is ungraded. It's for you to see how some git commands work and have fun with it! Follow the set of instructions below:

- Save this file.
- On the root of your Prep Course repository, run the `git status` command.
- Check the output, you should see that the Exercise Notebook file is modified.
- Run the `git add` command to add **just** the Exercise Notebook file.
- Run the `git commit` command with the option that allows you to write the commit message in the command. Write the following message: "I've harnessed the power of git!".

Congratulations! You've just finished the first steps of your journey into the wonderful world of version control! These skills will accompany you for the rest of the Prep Course (and maybe life ;) ), so make good use of them! 

More Git fun follows in SLU06.

# Submit your work!

To submit your work, follow the steps in [section 1.3](https://github.com/LDSSA/ds-prep-course-2025/blob/main/docs/weekly-workflow.md#13-work-on-the-learning-units) of the weekly workflow, starting from the step 3 "Grading the Exercise Notebook"!