# SLU06 - Git Intermediate

I hope you enjoyed the learning notebook! Like SLU03, the Exercise Notebook is multiple choice so it's really important that you have a good grasp of everything before starting it. 

**Good luck!**

```python
import hashlib
import requests
```

## Exercise 1: Git remotes

What is the main purpose of a remote repository?
- **A**: To serve as a backup to your local files.
- **B**: Tracking and committing changes to your files.
- **C**: To allow teammates to collaborate on a project.
- **D**: See who committed what changes to a project and hunt them down.

<!-- grade_id: cell-14d1430f3cf9833a, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_1.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_1 =
```

<!-- grade_id: cell-fccabbb7b3b73a9a, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_1, str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_1.upper(), encoding='utf8')).hexdigest() == \
'6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d', 'Not correct, try again.'
```

## Exercise 2: Git remotes

What is the function of `git push`?

- **A**: Updates the remote repository with the changes made in the local repository.
- **B**: Push all changes to the remote branch.
- **C**: Push all changes to the main branch.
- **D**: Updates the local repository with the changes made in the remote repository.

<!-- grade_id: cell-349ecbd85ce6985f, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_2.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_2 =
```

<!-- grade_id: cell-dbaa29386c745914, points: 2, grade: true, solution: false -->
```python
assert isinstance(exercise_2,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_2.upper(), encoding='utf8')).hexdigest() == \
'559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd', 'Not correct.'
```

## Exercise 3: Git remotes

What is the function of `git clone`?

- **A**: Creates a new Git repository on your local machine.
- **B**: Creates a copy of an existing Git repository in your remote repository.
- **C**: Creates a copy of the files in a Git repository, without any history.
- **D**: Creates a copy of an existing Git repository on your local machine.

<!-- grade_id: cell-548374c265341b10, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_3.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_3 =
```

<!-- grade_id: cell-9bb7f9bb7974e4d7, points: 4, grade: true, solution: false -->
```python
assert isinstance(exercise_3,str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_3.upper(), encoding='utf8')).hexdigest() == \
'3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43', 'Not correct.'
```

## Exercise 4: Git remotes

Did you review how to create and clone a repository [here](https://github.com/LDSSA/ds-prep-course-2025/blob/main/docs/github.md#32-setup-your-workspace-repository)? Good!

Now go to [Github](https://github.com/) and create a new repository called *SLU06* and then clone it to your machine.

**Important:** Make sure you set your repository to public!

It's the same thing you did when setting up the *ds-prep-workspace* repository, just with a different name. 

After creating the repository on GitHub, clone it to your local machine following the instructions [here](https://github.com/LDSSA/ds-prep-course-2025/blob/main/docs/github.md#33-clone-your-workspace-repository).

That's it for this exercise. We will use this new repository for the remaining exercises.

<!-- grade_id: cell-f7886da512b156a0, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_4 the URL of your SLU06 repository on GitHub, as a string.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_4 = ""
```

<!-- grade_id: cell-b99b518a1ed4e6d9, points: 4, grade: true, solution: false -->
```python
assert isinstance(exercise_4, str), "The variable should be a string."
request = requests.get(exercise_4)
assert request.status_code == 200, "Can't find the repository! Have you set it as public?"
```

## Exercise 5: Pushing changes into a remote repository

Follow this workflow:

1. In your **local SLU06 repository** on the main branch, create a `groceries.txt` file and add some grocery items to it.
2. Stage and commit your changes.
3. You just forgot the butter! Add it to your grocery list (`groceries.txt`).
4. Stage and commit your changes (again!).
5. Verify, by using `git status` and `git log`, that you have 2 commits in the local repository that are not in the remote repository.
6. Push the changes to the remote repository.
7. Verify, by using `git status` and `git log`, that your local and remote repositories are in sync.

If you like, repeat the workflow with more files. You can also create new branches. The idea is to get used to the commands. You will use them a lot, so start getting that muscle memory now!

## Exercise 6: Pulling changes from remote

Now try to do the same workflow we did in the learning notebook:

1. Use your browser to go to your SLU06 repository on Github (the one you created in exercise 4).
2. Edit your `groceries.txt` file on the GitHub website - Your mum wants eggs to make a pudding! ;)
3. Commit your changes on the website.
4. On your local machine, check that `git status` and `git log` don't show any differences. *If you feel like it, disconnect your computer from the network and check that those commands keep working :)*
5. Use `git remote` to check if the repositories are synced.
6. Pull the changes from the remote into your local repository.

## Exercise 7: Stash Changes

Follow this workflow in your local *SLU06* repository:

1. Create a new file named `winelist.txt`.
2. Add a few wines you'd like to drink (*yey quarantine for making day drinking a thing!*).
3. Stash the changes. Remember that stash does not save untracked files by default, so you need to use the appropriate option.
4. Add a new item to the `groceries.txt` file (you always need a massive amount of toilet paper in a global crisis).
5. Commit the changes.
6. Pop the changes you have in the stash.
7. If you're feeling like it, add a nice bottle of Moscatel to your winelist. When you're ready, commit.
8. Check, using whatever method you like, that the local repository is 2 commits ahead of the remote repository.
9. Push the commits to the remote repository.

## Exercise 8: Create and push branches

In you local *SLU06* repository, follow this workflow:

1. Create a new branch based on `main`. Name it "ubereats".
2. In the new branch create a new file `restaurants_to_try.txt` and add some content to it.
3. Commit the changes and push the branch to the remote.
4. Check on GitHub that the branch is there.
5. Create another branch based on `main`. Name it "movies".
6. In the new branch, create a file called `movies_to_watch` and add some content to it.
7. Commit the changes and push the branch to the remote.
8. Check on GitHub that the branch is there.

## Exercise 9: Merging branches

Follow this workflow in your local *SLU06* repository. I'm leaving out a lot of step-by-step instructions in this exercise and the next one. Remember that the learning notebook is your friend.

1. Merge the *ubereats* branch you created in exercise 8 into *main*. Verify that it was a *fast-forward* merge.
2. Merge the *movies* branch to the main. This one should be a *3-way merge*.
3. Push main to the remote.
4. Bonus: Do a Google search on how to delete a branch and delete the branches you created, both in the local and remote repositories.

**Well, the moment of truth has arrived...**

<img width="400" src="media/michael-scott.jpeg"/>

Were you able to do exercises 5-9?

- **A**: No, I've been stuck on this for more than an hour! I need to ask for help on Slack!
- **B**: Yes, I did it!

Really, take your time :) enjoy, learn, discover all these new things... but don't get frustrated!!! If you have any questions or are stuck somewhere, go to the Week04 slack channel and tell us what difficulties you are having, we will be happy to help you. <333

<!-- grade_id: cell-e8b4366001176c3a, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_5_9.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_5_9 =
```

<!-- grade_id: cell-d479dadee005a917, points: 0, grade: true, solution: false -->
```python
assert isinstance(exercise_5_9, str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_5_9.upper(), encoding='utf8')).hexdigest() == \
'df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c', "Please ask for help on Slack!"
```

## Exercise 10: Collaboration on Git

From the following statements about pull requests and branches, select the one that is false.

- **A**: A branch is just a separate version of the code.
- **B**: A pull request is when someone who has been working on a project branch wants to merge their work into another branch and requests a code review before doing it.  
- **C**: You can only do pull requests to the main branch, but you can branch from any other branch, including the main.
- **D**: A pull request lets the project maintainer know that an update is ready to be integrated.

<!-- grade_id: cell-4fc49f4aad340d87, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_10.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_10 =
```

<!-- grade_id: cell-813abdc5b03ac7ce, points: 4, grade: true, solution: false -->
```python
assert isinstance(exercise_10, str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_10.upper(), encoding='utf8')).hexdigest() == \
'6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d', "Not correct."
```

## Exercise 11: Merge conflicts

From the following statements about merge conflicts, select the one that is false.

- **A**: You can abort the merge and discard the alterations that are in conflict.
- **B**: You can edit the files to fix the conflicting changes and then add the resolved files by running `git add`, followed by `git commit`. 
- **C**: You can choose what version you want to keep of the two conflicting branches.
- **D**: After a conflict is created, you cannot go back to your previous state.

<!-- grade_id: cell-bb2837e739fc4c32, points: 0, grade: false, solution: true -->
```python
# Assign the correct option letter to the variable exercise_11.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_11 =
```

<!-- grade_id: cell-9e36d4171ea208f0, points: 4, grade: true, solution: false -->
```python
assert isinstance(exercise_11, str), "The variable should be a string."
assert hashlib.sha256(bytes(exercise_11.upper(), encoding='utf8')).hexdigest() == \
'3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43', "Not correct."
```

Now you're ready to make the most out of Git! You'll be using it throughout the rest of the Prep Course so it's important that you two become good friends!

Keep the git workflow in your heart, and use it whenever you find a typo in the learning materials! I'll give you a golden star for it! :D

# Submit your work!

To submit your work, follow the instructions in step 3 in [section 1.3](https://github.com/LDSSA/ds-prep-course-2025/blob/main/docs/weekly-workflow.md) of the weekfly workflow.